```python
# Initialize Otter
import otter
grader = otter.Notebook("02_errors.ipynb")
```

<center><h2> Assignment 02: Errors </h2></center>

Note: Please upload your solution as an ipynb file to the Gradescope page.

<justify> The purpose of this assignment is to develop your skills in computing absolute and relative errors. Use 64 bit floating point representation unless otherwise specified (this is the default in Python).</justify>

**Question 1:** <br>
if $\vert x \vert < 1$ it is known that:

<center>$\frac{1}{1 - x} = 1 + x + x^2 + x^3 + \dots$<br></center>
For this series, compute for $x = -0.1$ the absolute and relative errors for 3, 5 and 7 terms in the series.

```python
q1_abs_err = [] #store the answers in order of absolute errors for 3, 5, and 7 terms
q1_rel_err = [] #store the answers in order of relative errors for 3, 5, and 7 terms
terms = [3, 5, 7]
for i in range(3):
    print(f'abs error for {terms[i]} terms: {q1_abs_err[i]}')
    print(f'rel error for {terms[i]} terms: {q1_rel_err[i]}')
```

```python
grader.check("q1")
```

**Question 2:** <br>
Evaluate $e^{-5}$ using two approaches:
		\begin{align*}
			e^{-x} &  = 1 - x + \frac{x^2}{2!} - \frac{x^3}{3!} + \dots \\
			\mathrm{and}\\
			e^{-x} = \frac{1}{e^x} & = \frac{1}{1 + x + \frac{x^2}{2!} + \frac{x^3}{3!} + \dots} 
		\end{align*}
		compare with the true value of $6.737947 \times 10^{-3}$ for $x = 5$. Use 20 terms to evaluate each series and compute the absolute and relative errors.
		

```python
# store answers in the lists
q2_abs_err = [] 
q2_rel_err = [] 
# print approach 1 errors
print(f'abs_err for approach 1: {q2_abs_err[0]}')
print(f'rel_err for approach 1: {q2_rel_err[0]}')
# print approach 2 errors
print(f'abs_err for approach 2: {q2_abs_err[1]}')
print(f'rel_err for approach 2: {q2_rel_err[1]}')
```

```python
grader.check("q2")
```

**Question 3:** <br>
The function $\cos(x)$ can be approximated using the Taylor series expansion as follows:

\begin{align*}
        \cos(x) = 1 - \frac{x^2}{2!} + \frac{x^4}{4!} - \frac{x^6}{6!}+ \dots = \sum_{n = 0}^{\infty}\frac{(-1)^n}{(2n)!}x^{2n}
\end{align*}

(a) Calculate the absolute and relative error associated with the Taylor series expansion of $\cos(\pi/4)$ using:<br>
- i. one term in the expansion,<br>
- ii. two terms in the expansion, and<br>
- iii. three terms in the expansion.<br>

(b) Compute the number of terms required to have a solution that is accurate for 10 significant figures.

```python
# store answers in the lists
q3_abs_err = []
q3_rel_err = []
# number of terms for 10 significant figures
num_terms = ...

for i in range(3):
    print(f'abs error for {i+1} term: {q3_abs_err[i]}')
    print(f'rel error for {i+1} term: {q3_rel_err[i]}')
print(f'It requires {num_terms} terms to have a solution that is accurate for 10 significant figures')
```

```python
grader.check("q3")
```

**Question 4:** <br>
The infinite series
	\begin{equation*}
		f(n) = \sum_{i = 1}^{n} \frac{1}{i^4}
	\end{equation*}
converges on a value of $f(n) = \pi^4/90$ as $n$ approaches infinity.
	Write a program in single precision to calculate $f(n)$ for $n = 10,000$
	by computing the sum from $i = 1$ to $10,000$. Using 16, 32 and 64 bit floating point representations, compute the absolute errors. 
	
*Hint*: use *np.float16(x)* to represent a variable **x** as a 16-bit floating point number. Use the same representation for both the variable used for the summation and inside the iteration loop.
	

```python
q4_abs_err = [] # store answers in order of the following results

#use np.float16() inside iteration
sums = np.float16(0)
for ...:
    sums += np.float16(...)
q4_abs_err.append(...)

#use np.float32() inside iteration 
sums = np.float32(0)

#use np.float64() inside iteration
sums = np.float64(0)

#use np.float16() for the final summation
sums = 0

#use np.float32() for the final summation
sums = 0

#use np.float64() for the final summation
sums = 0

```

```python
grader.check("q4")
```

**Question 5:** <br>
Let $x^* = x \pm \varepsilon_x$ and $y^* =  y \pm \varepsilon_y$ are the absolute errors associated with the variables. Compute the absolute error accumulation for the subtraction operation $x^* - y^*$. (Refer to error accumulation section in the handout).

```python
x = 9.99
y = 9.98
ϵx = 0.005
ϵy = 0.005
# absolute error accumulation for the subtraction operation
ϵ = ...
```

```python
grader.check("q5")
```

**Question 6:** <br>
Let $x^* = x (1 \pm \eta_x)$ and $y^* =  y  (1 \pm \eta_y)$ are the relative errors associated with the variables. What is the relative error of multiplication ($x^*y^*$). Assume $\eta_x\eta_y$ is negligible as $\eta_x$ and $\eta_y$ are small enough.


```python
ηx = ...
ηy = ...
# the relative error of multiplication
η = ...
```

```python
grader.check("q6")
```

**Question 7:** <br>
The "divide and average" method is an old-time method for approximating the square root of any positive number $x$, and can be formulated as
	\begin{equation*}
		g = \frac{g + x / g}{2}
	\end{equation*}
	Write a well-structured function to implement this algorithm that computes the square root of a positive number up to 5 significant figures. Use ***math.sqrt(x)*** to get the exact value.

*Hint*: Square root of a number $x$ is $y$ such that $y*y = x$.
	The recipe for deducing the square root of a number $x$:
    <br>(a) Start with a guess, $g$. What could be a good initial guess?
    <br>(b) If $g$ is *close enough* to $\sqrt{x}$, stop and say $g$ is the answer. This is where the 5 signficant figures condition should be applied. 
    <br>(c) Otherwise make a *new guess* by averaging $g$ and $x/g$ as: $	g = \frac{g + x / g}{2}$
    <br>(d) Using a new guess, *repeat* the process until *close enough*. 

Check if your solution works for $x = 16$ and $x = \sqrt{2}$. Compute the relative error in both cases.

Example calculation:<br>

|$g$|$g*g$|$x/g$|$\frac{g + x/g}{2}$|
|:-|:-|:-|:-|
|3|9|16/3|4.167|
|4.167|14.36|3.837|4.0035|
|4.0035|16.0277|3.997|4.000002|

```python
q7_rel_err = [] # store the relative errors for x=16 and x=2**0.5
# x = 16
# x = math.sqrt(2)
```

```python
grader.check("q7")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert round(q1_abs_err[0] - 0.0009090909090909705, 9) == 0
```
```
>>> assert round(q1_rel_err[0] - 0.0010000000000000675, 9) == 0
```
```
>>> assert round(q1_abs_err[1] - 9.090909090958554e-06, 9) == 0
```
```
>>> assert round(q1_rel_err[1] - 1.000000000005441e-05, 9) == 0
```
```
>>> assert round(q1_abs_err[2] - 9.09090910328203e-08, 10) == 0
```
```
>>> assert round(q1_rel_err[2] - 1.0000000013610233e-07, 9) == 0
```


Tests for q2 (points: None)
```
>>> assert round(q2_abs_err[0] - 3.160594578443046e-05, 9) == 0
```
```
>>> assert round(q2_rel_err[0] - 0.00469073825965542, 9) == 0
```
```
>>> assert round(q2_abs_err[1] - 2.3251170901822382e-09, 11) == 0
```
```
>>> assert round(q2_rel_err[1] - 3.4507797259049946e-07, 9) == 0
```


Tests for q3 (points: None)
```
>>> assert round(q3_abs_err[0] - 0.2928932188134524, 9) == 0
```
```
>>> assert round(q3_rel_err[0] - 0.414213562373095, 9) == 0
```
```
>>> assert round(q3_abs_err[1] - 0.01553191872059001, 9) == 0
```
```
>>> assert round(q3_rel_err[1] - 0.021965450104334962, 9) == 0
```
```
>>> assert round(q3_abs_err[2] - 0.000322425523225478, 9) == 0
```
```
>>> assert round(q3_rel_err[2] - 0.00045597854780071235, 9) == 0
```
```
>>> assert num_terms == 7
```


Tests for q4 (points: None)
```
>>> assert round(q4_abs_err[0] - 0.000291983711137922, 9) == 0
```
```
>>> assert round(q4_abs_err[1] - 1.1130446340157363e-06, 9) == 0
```
```
>>> assert round(q4_abs_err[2] - 2.7688962234151404e-13, 15) == 0
```
```
>>> assert round(q4_abs_err[3] - 0.000291983711137922, 9) == 0
```
```
>>> assert round(q4_abs_err[4] - 4.016102805870503e-08, 10) == 0
```
```
>>> assert round(q4_abs_err[5] - 2.7688962234151404e-13, 15) == 0
```


Tests for q5 (points: None)
```
>>> assert ϵ == 0.01
```


Tests for q6 (points: None)
```
>>> assert round(η - 0.0010015025045085166, 9) == 0
```


Tests for q7 (points: None)
```
>>> assert q7_rel_err[0] < 1e-5/4
```
```
>>> assert q7_rel_err[0] < 1e-5/math.sqrt(math.sqrt(2))
```
