```python
# Initialize Otter
import otter
grader = otter.Notebook("05_roots.ipynb")
```

## Exercise 05.1

Write root finding functions using (bisection method, Newton Raphson and fixed point iteration) to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

Compute the roots of the following Polynomial function $ x^3 - 6x^2 + 4x + 12$  using:
* Bisection method
* Newton Raphson iterations
* Fixed point iteration



*HINT* <br>
raise ValueError for invalid initial guess(x0,x1) of bisection method<br>
raise RuntimeError for max iteration exceeded

```python
def fn(x):
    ...

# biscection
def bisection(f, x0, x1, tol, max_it):
    ...

# Newton-Raphson
def dfn(x):
    ...

def newton_raphson(f, df, x0, tol, max_it):
    ...

    
bisection(fn,4, 6,1e-6,100)
```

```python
# bisection
# Test with max_it = 30
x, error, num_it = bisection(fn, x0=3, x1=6, tol=1.0e-6, max_it=30)
assert round(x-4.5340702, 5) == 0
assert num_it > 20
```

```python
# Newton-Raphson
# Test with max_it = 30
x, error, num_it = newton_raphson(fn, dfn, x0=3, tol=1.0e-6, max_it=30)
assert round(x-2.51730405, 5) == 0
assert num_it > 4
```

```python
grader.check("q1")
```

### Fixed-point iteration

Usually a formula for finding the root of an equation can be found by rearranging $f(x) = 0$ to be: $x = g(x)$ and then using the computation formula: $$ x_{i+1} = g(x_i)$$ to solve for succesively more accurate approximations of the root. 

Consider: $f(x) = x^2 - 4\sin(x) =0 $ can be rearranged as $x = g(x) = 4\frac{\sin(x)}{x}$. So the computational formula is $$ x_{i+1} = g(x_i) = 4\frac{\sin(x_i)}{x_i}$$

We could solve the equations for $x$ stating at an initial guess of $x = x_0$. The relative approximation error is computed as:

$$\eta = \left|\frac{x_{i+1} - x_i}{x_{i+1}}\right| < \varepsilon$$.


```python
def fixed_point(f, x0, tol, max_it):
    ...
# find 1 of 3 roots with fixed point method
# you may try different x0 and max iterations to find the root
x, error, num_it = fixed_point(fn, x0=..., tol=1.0e-6, max_it=...)
root = ... # store the root in this variable
```

```python
import numpy as np
assert ((np.roots([1,-6,4,12])-root).round(4)==0).any()
```

```python
grader.check("q2")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> # bisection
>>> # Test with max_it = 30
>>> x, error, num_it = bisection(fn, x0=3, x1=6, tol=1.0e-6, max_it=30)
>>> assert round(x-4.5340702, 5) == 0
>>> assert num_it > 20
```
```
>>> # Newton-Raphson
>>> # Test with max_it = 30
>>> x, error, num_it = newton_raphson(fn, dfn, x0=3, tol=1.0e-6, max_it=30)
>>> assert round(x-2.51730405, 5) == 0
>>> assert num_it > 4
```


Tests for q2 (points: None)
```
>>> import numpy as np
>>> assert ((np.roots([1,-6,4,12])-root).round(4)==0).any()
```
