```python
# Initialize Otter
import otter
grader = otter.Notebook("04_data_structures.ipynb")
```

## Exercise 04.1 (selecting and passing data structures)

Given the coordinates of the vertices of a triangle, $(x_0, y_0)$, $(x_1, y_1)$ and $(x_2, y_2)$, the area $A$ of the triangle is given by:
$$
A = \left| \frac{x_0(y_1  - y_2) + x_1(y_2 - y_0) + x_2(y_0 - y_1)}{2} \right|
$$
There is an old function named `area` that computes the area of a triangle given the coordinates of the vertices as six arguments ($x$ and $y$ components of each vertex). The order of the function arguments must be (`x0, y0, x1, y1, x2, y2)`. With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

You have been hired as a Python consultant to rewrite the function for computing the area of a triangle.

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.<br>
**HINT:** <li>if the number of input coordinates is not 3, print "incorrect number of coordinates"</li>
<li>if the calculated area is 0, print "invalid coordinates for a triangle"</li>

```python
# input format
# coord = [(x0,y0),(x1,y1),(x2,y2)]
# area_triangle(coord)
# YOUR CODE HERE
def area_triangle(...):
    ...
```

```python
a = [(1,1),(2,5),(3,1)]
area_triangle(a)
```

```python
a = [(1,1),(2,5)]
area_triangle(a)
```

```python
a = [(1,1),(2,5),(3,1),(2,0)]
area_triangle(a)
```

```python
a = [(1,1),(2,5),(2,5)]
area_triangle(a)
```

```python
grader.check("q1")
```

## Exercise 04.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

```python
def polygon_area
```

```python
vertices = [(1,1),(2,5),(3,1)]
polygon_area(vertices)
```

```python
vertices=[(1,1),(2,5),(5,5),(6,1)]
polygon_area(vertices)
```

```python
grader.check("q2")
```

## Exercise 04.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
def sum_vector(...):
    ...
```

```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
sum_vector(a,b)
```

```python
grader.check("q3")
```

## Exercise 04.4 (dictionaries)

Create a dictionary that maps build names (the key) to building abbreviations for at least 5 buildings
(you can find abbreviations at https://utdirect.utexas.edu/apps/campus/buildings/nlogon/facilities/UTM/).
From the dictionary, produce and print

1. A dictionary from building abbreviation to name; and
1. A list of building abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps building names (the key) to dictionaries of:

- Building abbreviation
- Year occupied
- Zones
 
for at least 5 building. Take the data from https://utdirect.utexas.edu/apps/campus/buildings/nlogon/facilities/UTM/. Using this dictionary, 

1. Find the oldest occupied building, and print the zone and the abbreviation for this building
2. Find the oldest building in a given zone.

```python
# find the abbreviations for the buildings
buildings = ["ATHLETIC FIELDS PAVILION (EASTSIDE)", 
            "ALMETRIS DUREN RESIDENCE HALL",
            "ATHLETIC FIELDS PAVILION (REHAB)",
            "ATHLETIC FIELDS PAVILION",
            "ANNA HISS GYMNASIUM"]
# make a dictionary that maps building names (key) to building abbreviations (value)
building_abv = ...

# make a dictionary that maps building abbreviations (key) to building names (value)
# Hint: use building_abv to create the this dictionary. Don't manually type keys and values...
abv_building = ...

# make a list of building abbreviations and sort them into alphabetical order
# Hint: use the dictionary functions .values() or .keys() to list the values or keys, and then use the function 'sorted' to sort the list.
sorted_abv = ...

# make a dictionary that maps building names to the years occupied
building_year = ...

# make a dictionary that maps building names to zones
building_zone = ...

# use the dictionaries above to find the oldest building in the dictionary
old_building = ...
old_zone = ...
old_abv = ...

# use the dictionaries above to find the oldest building in zone 3
old_building2 = ...
```

```python
sorted_abv
```

```python
old_building
```

```python
old_abv
```

```python
old_zone
```

```python
old_building2
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> a = [(1,1),(2,5),(3,1)]
>>> area_triangle(a)
4.0
```
```
>>> a = [(1,1),(2,5)]
>>> area_triangle(a)
incorrect number of coordinates
```
```
>>> a = [(1,1),(2,5),(3,1),(2,0)]
>>> area_triangle(a)
incorrect number of coordinates
```
```
>>> a = [(1,1),(2,5),(2,5)]
>>> area_triangle(a)
invalid coordinates for a triangle
```


Tests for q2 (points: None)
```
>>> vertices = [(1,1),(2,5),(3,1)]
>>> polygon_area(vertices)
4.0
```
```
>>> vertices=[(1,1),(2,5),(5,5),(6,1)]
>>> polygon_area(vertices)
16.0
```


Tests for q3 (points: None)
```
>>> a = [0, 4.3, -5, 7]
>>> b = [-2, 7, -15, 1]
>>> sum_vector(a,b)
[-2, 11.3, -20, 8]
```


Tests for q4 (points: None)
```
>>> sorted_abv
['ADH', 'AF1', 'AF2', 'AFP', 'AHG']
```
```
>>> old_building
'ANNA HISS GYMNASIUM'
```
```
>>> old_abv
'AHG'
```
```
>>> old_zone
1
```
```
>>> old_building2
'ATHLETIC FIELDS PAVILION'
```
