## Exercise 01.0 (if-else)

Consider the following assessment criteria which map a score out of 100 to an assessment grade for CE311K:

| Grade            | Raw score  (/100)      |
| ---------------- | ---------------------- |
| A                | $\ge 90$               |
| A-               | $\ge 87$ and $< 90$    |
| B+               | $\ge 84$ and $< 87$    |
| B                | $\ge 80$ and $< 84$    |
| B-               | $\ge 77$ and $< 80$    |
| C+               | $\ge 74$ and $< 77$    |
| C                | $\ge 70$ and $< 74$    |
| C-               | $\ge 67$ and $< 70$    |
| D+               | $\ge 64$ and $< 67$    |
| D                | $\ge 60$ and $< 64$    |
| F                | $< 60$                 |

Write a program that, given an a score, prints the appropriate grade. Print an error message if the input score is greater than 100 or less than zero.

<!-- grade_id: cell-71121f470be22e4a, points: 0, grade: false, solution: true -->
```python
scores = [71, 90, 59, 60, -5, 100, 101]

# Output grade as string
grade = []

# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-1-5ba36e73ab6b> in <module>()
      5 
      6 # YOUR CODE HERE
----> 7 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-3a0d0a59613d6b3f, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 3.0
x1 = 6.0
x_mid = (x0 + x1)/2


# Print the number of iterations, final x and final f value
print(n, x_mid, f)
```

<!-- grade_id: cell-421824bbd4e61fae, points: 0, grade: true, solution: false -->
```python
assert round(x_mid - 4.534149169921875, 10) == 0.0
assert abs(f) < 0.0009
```

<!-- grade_id: cell-1747a3fb1dd1de48, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 3.0
x1 = 6.0

tol = 1.0e-6
error = tol + 1.0 #Initial error just larger than tolerance


print(counter, x_mid, f, error)
```

<!-- grade_id: cell-ddff32f44df88388, points: 0, grade: true, solution: false -->
```python
assert counter == 23
assert abs(f) < 1.0e-6
```

## Exercise 01.2 (series expansion)

The power series expansion for the sine function is: 

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

(See mathematics data book for a less compact version; this compact version is preferred here as it is simpler to program.)

1. Using a `for` statement, approximate $\sin(3\pi/2)$ using 15 terms in the series expansion and report the absolute error.

1. Using a `for` statement, compute how many terms in the series are required to approximate $\sin(3\pi/2)$ to within $1 \times 10^{-8}$. 

Store the absolute value of the error in the variable `error`.

*Note:* Calculators and computers use iterative or series expansions to compute trigonometric functions, similar to the one above (although they use more efficient formulations than the above series).

### Hints

To compute the factorial and to get a good approximation of $\pi$, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
pi = math.pi
```
You only need '`import math`' once at the top of your program. Standard modules, like `math`, will be explained in a later. If you want to test for angles for which sine is not simple, you can use 
```python
a = 1.3
s = math.sin(a)
```    
to get an accurate computation of sine to check the error.

#### (1) Using a `for` loop

<!-- grade_id: cell-d6df4b428f969a5d, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.sin and math.factorial
import math

# Value at which to approximate sine
x = 1.5*math.pi

# Initialise approximation of sine
approx_sin = 0.0

# YOUR CODE HERE
raise NotImplementedError()
    
print("The error is:")
print(error)
```

<!-- grade_id: cell-8b269b717e92414d, points: 0, grade: true, solution: false -->
```python
assert error < 1.0e-12
```

#### (2) Using a `for` loop but defined tolerance

*Remember to guard against infinite loops.*

<!-- grade_id: cell-03df25d83df3674e, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.sin and math.factorial
import math

# Value at which to approximate sine
x = 1.5*math.pi

# Tolerance and initial error (this just needs to be larger than tol)
tol = 1.0e-8
error = tol + 1.0

# Intialise approximation of sine
approx_sin = 0.0

# Initialise counter
n = 0

# Loop until error satisfies tolerance or a fixed number of iterations is reached
for i in range(1001):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
print("The error is:")
print(error)

print("Number of terms in series:")
print(n)
```

<!-- grade_id: cell-03df25d83df3674dfwe, points: 0, grade: true, solution: false -->
```python
assert error <= 1.0e-8
```