```python
# Initialize Otter
import otter
grader = otter.Notebook("02b_errors.ipynb")
```

## Exercise 02b.0

Compare the computed values of 

$$ 
d_0 = a \cdot b + a  \cdot  c
$$

and 

$$ 
d_1 = a \cdot (b + c)
$$

when $a = 100$, $b = 0.1$ and $c = 0.2$. Store $d_{0}$ in the variable `d0` and $d_{1}$ in the variable `d1`.


Try checking for equality, e.g. `print(d0 == d1)`. 

```python
# YOUR CODE HERE

```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

```python
grader.check("q1")
```

## Exercise 02b.1

For the polynomial 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
compute $f$ using: (i) the compact form $(x + y)^{6}$; and (ii) the expanded form for:

(a) $x = 10$ and $y = 10.1$

(b) $x = 10$ and $y = -10.1$

and compare the number of significant digits for which the answers are the same.
Store the answer for the compact version using the variable `f0`, and using the variable `f1` for the expanded version.

For case (b), compare the computed and analytical solutions and consider the relative error.
Which approach would you recommend for computing this expression?

_Type your answer here, replacing this text._

```python
# x = 10 and y = 10.1
x = 10.0
y = 10.1

# YOUR CODE HERE
raise NotImplementedError()
```

```python
import math
ans1 = [f0, f1]
assert math.isclose(ans1[0], 65944160.60120103, rel_tol=1e-10)
assert math.isclose(ans1[1], 65944160.601201, rel_tol=1e-10)
```

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
# x = 10 and y = -10.1
x = 10.0
y = -10.1

# YOUR CODE HERE
raise NotImplementedError()
```

```python
import math
ans2 = [f0, f1]
assert math.isclose(ans2[0], 1.0e-6, rel_tol=1e-10)
assert math.isclose(ans2[1], 1.0e-6, rel_tol=1e-2)
```

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
#Using 16 bit floating points with numpy for (a) x=10 and y=10.1
x = 10.0
y = 10.1

# YOUR CODE HERE
raise NotImplementedError()
```

```python
import math
ans3 = [f0, f1]
assert math.isclose(ans3[0], 65821226.05029293, rel_tol=1e-10)
assert math.isclose(ans3[1], 65974924.11685877, rel_tol=1e-10)
```

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
#Using 16 bit floating points with numpy for (b) x=10 and y=-10.1
x = 10.0
y = -10.1

# YOUR CODE HERE
raise NotImplementedError()
```

```python
import math
ans4 = [f0, f1]
assert math.isclose(ans4[0], 1.0974893029924715e-06, rel_tol=1e-10)
assert math.isclose(ans4[1], 1.0977964848279953e-06, rel_tol=1e-2)
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

## Exercise 02b.2

Consider the expression

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

When $x$ is very large, the denominator approaches zero, which can cause problems.

Try rephrasing the problem and eliminating the fraction by multiplying the numerator and denominator by $\sqrt{x^2 - 1} + x$ and evaluate the two versions of the expression when:

(a) $x = 1 \times 10^{7}$

(b) $x = 1 \times 10^{9}$ (You may get a Python error for this case. Why?)

Store the answer for the original version using the variable `f0`, and using the variable `f1` for the rephrased version.

#### (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f0,f1)
```

#### (b) $x = 1 \times 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f1)
```

#### Using 32 bit floating points compute (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f0,f1)
```

#### Using 16 bit floating points compute (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f0,f1)
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert d0 == 30.0
```
```
>>> assert d1 != 30.0
```
```
>>> assert d0 != d1
```


Tests for q2 (points: None)
```
>>> assert math.isclose(ans1[0], 65944160.60120103, rel_tol=1e-10)
>>> assert math.isclose(ans1[1], 65944160.601201, rel_tol=1e-10)
```
```
>>> assert math.isclose(ans2[0], 1.0e-6, rel_tol=1e-10)
>>> assert math.isclose(ans2[1], 1.0e-6, rel_tol=1e-2)
```
```
>>> assert math.isclose(ans3[0], 65821226.05029293, rel_tol=1e-10)
>>> assert math.isclose(ans3[1], 65974924.11685877, rel_tol=1e-10)
```
```
>>> assert math.isclose(ans4[0], 1.0974893029924715e-06, rel_tol=1e-10)
>>> assert math.isclose(ans4[1], 1.0977964848279953e-06, rel_tol=1e-2)
```
