## Exercise 00

Degrees Fahrenheit ($T_F$) are converted to degrees Celsius ($T_c$) using the formula:

$$
T_c = 5(T_f - 32)/9
$$

Write a program to convert 78 degrees Fahrenheit to degrees Celsius and print the result.
Write your program such that you can easily change the temperature in Fahrenheit that you are converting from.

Use the variable name `T_c` for the temperature in degrees Celsius.

<!-- grade_id: cell-ca2d6af7fabf2918, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-a7542313c8cdf4ed, points: 0, grade: true, solution: false -->
```python
assert round(T_c - 25.55555555555555, 10) == 0.0
```

## Exercise 00

Evaluate: 
$$
v0 = \frac{(14.8^2 + 6.5^2)}{3.8^2} + \frac{55}{\sqrt{2} + 14}
$$

```python

```

```python
assert round(v0 - 21.6630, 4) == 0.0
```

<!-- grade_id: cell-aaab53918bb9627d, points: 0, grade: false, solution: true -->
## Exercise 01
Evaluate: 
$$
v1 = 2.34 + \frac{1}{2} 2.7(5.9^2 - 2.4^2) + 9.8 \ln(51)
$$

> Hint: Use `import math` function for log and trignometric functions

```python

```

```python
assert round(v1 - 80.0894, 4) == 0.0
```

## Exercise 02
Evaluate: 
$$
v2 = \frac{\tan 64^o}{\cos^2 14^o} - \frac{3 * \sin 80^o}{\sqrt[3]{0.9}} + \frac{\cos{55^o}}{\sin 11^o}
$$

```python

```

```python
assert round(v2 - 2.1238, 4) == 0.0
```

## Exercise 03

Define the variable `x` as `2.34`, then evaluate: 
Evaluate: 

$$
v3 = \frac{e^{2x}}{\sqrt{14 + x^2 - x}}
$$

```python

```

```python
assert round(v3 - 26.0345, 4) == 0.0
```

## Exercise 04

Define the variables `a`, `b`, `c`, and `d` as: `a = 13`, `b=4.2`,
$c = 4b/a$, $d = (abc)/(a + b + c)$. The evaluated:

$$
v4 = \frac{\sqrt{a^2 + b^2}}{(d - c)} + \ln(\left|b - a + c -d\right|)
$$

```python

```

```python
assert round(v4 - 7.8410, 4) == 0.0
```

## Exercise 05

Measure the arc length of the cable of the parabolic segment BOC supporting the Golden Gate Bridge. 

The arc legnth `s` of a parabolic segment BOC is gent as:

![golden gate bridge](golden-gate-bridge.png)
    
$$
s = \frac{1}{2}\sqrt{b^2 + 16 a^2} + \frac{b^2}{8a}\ln\left(\frac{4a + \sqrt{b^2 + 16 a^2}}{b}\right)
$$

Given $a = 40 m$ and $b = 2000 m$

```python

```