Make NumPy available:

```python
import numpy as np
```

## Exercise 07.1 (indexing and timing)

Create two very long NumPy arrays `x` and `y` and sum the arrays using:

1. The NumPy addition syntax, `z = x + y`; and
2. A `for` loop that computes the sum entry-by-entry

Compare the time required for the two approaches for vectors of different lengths (use a very long vector for 
the timing). The values of the array entries are not important for this test. Use `%time` to report the time.

*Hint:* To loop over an array using indices, try a construction like:

```python
x = np.ones(10)
y = np.ones(len(x))
for i in range(len(x)):
    print(x[i]*y[i])
```

#### (1) Add two vectors using built-in addition operator and test using a random vector of size 10:

<!-- grade_id: cell-2c856e54f7c3340e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### (2) Add two vectors using own implementation:

<!-- grade_id: cell-f34614f9f0068cc4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Optional extension: just-in-time (JIT) compilation

You will see a large difference in the time required between your NumPy and 'plain' Python implementations. This is due to Python being an *interpreted* language as opposed to a *compiled* language. A way to speed up plain Python implementions is to convert the interpreted Python code into compiled code. A tool for doing this is [Numba](https://numba.pydata.org/).

Below is an example using Numba and JIT to accelerate a computation:

```python
!pip -q install numba 
import numba
import math

# Using an index operator using math.sin and return a vector
def compute_sine_native(x):


# Using an index operator using math.sin and return a vector
# The only difference is the addition of `@numba.jit` at the top
@numba.jit
def compute_sine_jit(x):
    
    

# Try using numpy sin value
def compute_sin_np(x):
    
    
    
    
x = np.ones(10000000)
%time z = compute_sine_native(x)
%time z = compute_sine_jit(x)
%time z = compute_sine_np(x)
```

**Task:** Test if Numba can be used to accelerate your implementation that uses indexing to sum two arrays, and by how much.

## Exercise 07.2 (member functions)

Anonymised scores (out of 60) for an examination are stored in a NumPy array. Write:

1. A function that takes a NumPy array of the raw scores and returns the scores as percentages, sorted from 
   lowest to highest (try using `scores.sort()`, where `scores` is a NumPy array holding the scores).
1. A function that returns the maximum, minimum and mean of the raw scores as a dictionary with the 
   keys '`min`', '`max`' and '`mean`'. Use the NumPy array functions `min()`, `max()` and `mean()` to do the 
   computation, e.g. `max = scores.max()`.  
   
1. Define a function `statistics_naive` without using the predefined NumPy array functions.
   
   Design your function for the min, max and mean to optionally exclude the highest and lowest scores from the 
   computation of the min, max and mean. 
   
   *Hint:* sort the array of scores and use array slicing to exclude
   the first and the last entries.

Use the scores 
```python
scores = np.array([58.0, 35.0, 24.0, 42, 7.8])
```
to test your functions.

<!-- grade_id: cell-169ebae60810c6be, points: 0, grade: false, solution: true -->
```python
def to_percentage_and_sort(scores):
    # YOUR CODE HERE
    raise NotImplementedError()

def statistics(scores, exclude=False):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def statistics_naive(scores, exclude=False):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-af0b6fd8a3cadb1a, points: 0, grade: true, solution: false -->
```python
scores = np.array([58.0, 35.0, 24.0, 42, 7.8])
assert np.isclose(to_percentage_and_sort(scores), [ 13.0, 40.0, 58.33333333,  70.0, 96.66666667]).all()

s0 = statistics(scores)
assert round(s0["min"] - 7.8, 10) == 0.0
assert round(s0["mean"] - 33.36, 10) == 0.0
assert round(s0["max"] - 58.0, 10) == 0.0

s1 = statistics(scores, True)
assert round(s1["min"] - 24.0, 10) == 0.0
assert round(s1["mean"] - 33.666666666666666667, 10) == 0.0
assert round(s1["max"] - 42.0, 10) == 0.0

s2 = statistics_naive(scores)
assert round(s2["min"] - 7.8, 10) == 0.0
assert round(s2["mean"] - 33.36, 10) == 0.0
assert round(s2["max"] - 58.0, 10) == 0.0

s3 = statistics_naive(scores, True)
assert round(s3["min"] - 24.0, 10) == 0.0
assert round(s3["mean"] - 33.666666666666666667, 10) == 0.0
assert round(s3["max"] - 42.0, 10) == 0.0
```

## Exercise 07.3 (angle between two vectors)

Compute the angle between two vectors $a$ and $b$. Using:
    $$\alpha = \cos^{-1}\left( \frac{a \cdot b}{|a| \cdot |b|} \right)$$
    
    
Where $|a|$ is the magnitude of the vector.


```python
# Your CODE HERE
raise NotImplementedError()
```

```python
a = np.array([3., 4., 0.])
b = np.array([4., 4., 2.])

angle = angle_between_vectors(a, b)
assert round(angle - 0.933333333, 10) == 0.0
```