## Exercise 03 functions

The objective of the lab is to familiarize with writing and calling custom functions in Python

### 03.0 interger multiple

Write a function called `multiple` to determine whether an integer `b` is a multiple of `a`. Your function should return `True` if `b` is a multiple of a, and `False` otherwise. 

```python
# b is a multiple of a

```

<!-- grade_id: cell-421824bbd4e61fae, points: 0, grade: true, solution: false -->
```python
assert multiple(15, 5) == True
assert multiple(17, 7) == False
```

### 03.1 Max of 3 numbers

Write a function called `max3` that takes 3 numbers, and returns the maximum of the 3.


```python
# Write a function `max3` that returns maximum of 3 numbers

```

```python
assert max3(-5.3, 2.5, 7.3) == 7.3
```

### 03.2 Stats from a list of numbers

Write a function called `stats` that takes a list as an argument and return the `max`, `min`, `mean`, and `sum` of the list. Please write custom code to compute the value of sum and mean rather than using the default functions for a Python list. You may use `len` to find the length of a list.

```python
# Stats of list

```

```python
# Check results of stats
numbers = [7, 5, 2, -3, 11, 17]

max, min, mean, sum = stats(numbers)
assert max == 17
assert min == -3
assert round(mean, 1) == 6.5
assert sum == 39
```

### 03.3 State of stress

The state of stress at an interior point in a solid is defined by three stress components $\sigma_{xx}$, $\sigma_{yy}$, and $\sigma_{xy}$. Given the three components, the principal stresses (maximum and minimum normal stresses) $\sigma_1$ and $\sigma_2$ can be computed from:

$$ \sigma_1 = \frac{\sigma_{xx} + \sigma_{yy}}{2} + \sqrt{\left(\frac{\sigma_{xx} - \sigma_{yy}}{2}\right)^2 + \sigma_{xy}^2}$$

$$ \sigma_2 = \frac{\sigma_{xx} + \sigma_{yy}}{2} - \sqrt{\left(\frac{\sigma_{xx} - \sigma_{yy}}{2}\right)^2 + \sigma_{xy}^2}$$

Write a function called `principal_stresses` that takes the 3 stress components as the arguments and returns `sigma_1` and `sigma_2` as the maximum and minimum principal stresses.

```python
# State of stress

```

```python
# Check state of stress
sigma_xx = 50
sigma_yy = -20
sigma_xy = 30
sigma_1, sigma_2 = principal_stresses(sigma_xx, sigma_yy, sigma_xy)
assert round(sigma_1, 1) == 61.1
assert round(sigma_2, 1) == -31.1
```

### 03.4 Recursive factorial

Write two functions called `factorial` and `recursive_factorial` to calculate the factorial of a number.
The `factorial` function may calculate the factorial using a `for` loop using a `range` function. The `recursive factorial` function must calculate the factorial, bu recursively calling the function `recursive_factorial`.

```python
# Define factorial(n)


# Define recursive_factorial(n)
```

```python
assert factorial(10) == 3628800
assert recursive_factorial(10) == 3628800
```