```python
# Initialize Otter
import otter
grader = otter.Notebook("03_functions.ipynb")
```

## Exercise 03 functions

The objective of the lab is to familiarize with writing and calling custom functions in Python

### 03.0 interger multiple

Write a function called `multiple` to determine whether an integer `b` is a multiple of `a`. Your function should return `True` if `b` is a multiple of a, and `False` otherwise. 

```python
def multiple(...):
```

```python
assert multiple(b=15, a=5) == True
assert multiple(b=17, a=7) == False
```

```python
grader.check("q1")
```

### 03.1 Max of 3 numbers

Write a function called `max3` that takes 3 numbers, and returns the maximum of the 3.


```python
def ...
```

```python
assert max3(-5.3, 2.5, 7.3) == 7.3
```

```python
grader.check("q2")
```

### 03.2 Stats from a list of numbers

Write a function called `stats` that takes a list as an argument and return the `max`, `min`, `mean`, and `sum` of the list. Please write custom code to compute the value of sum and mean rather than using the default functions for a Python list. You may use `len` to find the length of a list.

```python
# you are not allowed to use any other python module or functions to finish this question.
def ...
```

```python
# Check results of stats
numbers = [7, 5, 2, -3, 11, 17]

max, min, mean, sum = stats(numbers)
assert max == 17
assert min == -3
assert round(mean, 1) == 6.5
assert sum == 39
```

```python
grader.check("q3")
```

### 03.3 State of stress

The state of stress at an interior point in a solid is defined by three stress components $\sigma_{xx}$, $\sigma_{yy}$, and $\sigma_{xy}$. Given the three components, the principal stresses (maximum and minimum normal stresses) $\sigma_1$ and $\sigma_2$ can be computed from:

$$ \sigma_1 = \frac{\sigma_{xx} + \sigma_{yy}}{2} + \sqrt{\left(\frac{\sigma_{xx} - \sigma_{yy}}{2}\right)^2 + \sigma_{xy}^2}$$

$$ \sigma_2 = \frac{\sigma_{xx} + \sigma_{yy}}{2} - \sqrt{\left(\frac{\sigma_{xx} - \sigma_{yy}}{2}\right)^2 + \sigma_{xy}^2}$$

Write a function called `principal_stresses` that takes the 3 stress components as the arguments and returns `sigma_1` and `sigma_2` as the maximum and minimum principal stresses.

```python
def principal_stresses(sigma_xx, sigma_yy, sigma_xy):
    ...
    return sigma_1, sigma_2
```

```python
# Check state of stress
sigma_xx = 50
sigma_yy = -20
sigma_xy = 30
sigma_1, sigma_2 = principal_stresses(sigma_xx, sigma_yy, sigma_xy)
assert round(sigma_1, 1) == 61.1
assert round(sigma_2, 1) == -31.1
```

```python
grader.check("q4")
```

### 03.4 Function as an argument

Write a function `is_fn_positive` that takes another function as an argument and a range and an increment, to check if a given function is always positive in the given range.

```python

def is_fn_positive(fn, min, max):
    ...
```

```python
import numpy as np
assert is_fn_positive(np.sin, -1, 1) == False
assert is_fn_positive(np.sin, 0.1, 1) == True
assert is_fn_positive(np.cos, -1, 1) == True
```

```python
grader.check("q5")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert multiple(b=15, a=5) == True
>>> assert multiple(b=17, a=7) == False
```


Tests for q2 (points: None)
```
>>> assert max3(-5.3, 2.5, 7.3) == 7.3
```


Tests for q3 (points: None)
```
>>> # Check results of stats
>>> numbers = [7, 5, 2, -3, 11, 17]
>>> 
>>> max, min, mean, sum = stats(numbers)
>>> assert max == 17
>>> assert min == -3
>>> assert round(mean, 1) == 6.5
>>> assert sum == 39
```


Tests for q4 (points: None)
```
>>> # Check state of stress
>>> sigma_xx = 50
>>> sigma_yy = -20
>>> sigma_xy = 30
>>> sigma_1, sigma_2 = principal_stresses(sigma_xx, sigma_yy, sigma_xy)
>>> assert round(sigma_1, 1) == 61.1
>>> assert round(sigma_2, 1) == -31.1
```


Tests for q5 (points: None)
```
>>> import numpy as np
>>> assert is_fn_positive(np.sin, -1, 1) == False
>>> assert is_fn_positive(np.sin, 0.1, 1) == True
>>> assert is_fn_positive(np.cos, -1, 1) == True
```
