```python
# Initialize Otter
import otter
grader = otter.Notebook("01_iteration_control_flow.ipynb")
```

## Exercise 01.0 (if-else)

Consider the following assessment criteria which map a score out of 100 to an assessment grade for CE311K:

| Grade            | Raw score  (/100)      |
| ---------------- | ---------------------- |
| A                | $\ge 90$               |
| A-               | $\ge 87$ and $< 90$    |
| B+               | $\ge 84$ and $< 87$    |
| B                | $\ge 80$ and $< 84$    |
| B-               | $\ge 77$ and $< 80$    |
| C+               | $\ge 74$ and $< 77$    |
| C                | $\ge 70$ and $< 74$    |
| C-               | $\ge 67$ and $< 70$    |
| D+               | $\ge 64$ and $< 67$    |
| D                | $\ge 60$ and $< 64$    |
| F                | $< 60$                 |

Step 1: Write a program that finds the grade for a score of 76 and print out the grade.

_Type your answer here, replacing this text._

```python
score = 76

# Write your code here 

```

```python
grade_check = 'C+'
assert grade == grade_check
```

```python
scores = [79, 75, 59, 60, -5, 100, 101]

# Output grades as a string list
grades = []

# Write your code here

```

```python
grades_check = ['B-', 'C+', 'F', 'D', 'Score error', 'A', 'Score error']
assert grades == grades_check
```

```python
grader.check("q0")
```

## Exercise 01.1 (for and if-else)

Consider a list of numbers shown below. Classify the list into the following separate lists allowing for a number to appear on multiple lists (e.g., 9 could be on odd and squares): even numbers, odd numbers, and squares.

`numbers = [0, 2, 5, 7, 9, 11, 4, 6, 25, 16, 39, 43]`

```python
numbers = [0, 2, 5, 7, 9, 11, 4, 6, 25, 16, 39, 43]

# Create empty list of sequences
even = []
odd = []
squares = []
primes = []
# Write your code here

```

```python
assert even == [0, 2, 4, 6, 16]
assert odd == [5, 7, 9, 11, 25, 39, 43]
assert primes == [2, 5, 7, 11, 43]
assert squares == [0, 9, 4, 25, 16]
```

```python
grader.check("q1")
```

## Exercise 01.2: Cartesian to polar coordinates

Convert the x and y coordinates of a point in a plane to polar coordinates (r, φ) in which φ is measured from
(-π, π) with φ = 0 defined as the x-axis. To function correctly, the calculations must appropriately evaluate which quadrant the point is in and it must prevent division by zero, otherwise the program will experience an error. The expressions required to compute the appropriate φ are:

$r = \sqrt(x^2 + y^2)$

$\theta = \tan^{-1}(y/x) \quad  \quad \quad$  if $ x > 0 $

$\theta = \tan^{-1}(y/x) + \pi \quad$ if $ x < 0 $ and $y \ge 0$

$\theta = \tan^{-1}(y/x) - \pi \quad$ if $ x < 0 $ and $y < 0$

$\theta = \pi/2  \quad \quad  \quad \quad$ if $ x = 0 $ and $y > 0$

$\theta = -\pi/2 \quad \quad  \quad \quad$ if $ x = 0 $ and $y < 0$

$\theta = 0 \quad \quad \quad  \quad$ if $ x = 0 $ and $y = 0$

![Cartesian to polar coordinates](https://upload.wikimedia.org/wikipedia/commons/7/78/Polar_to_cartesian.svg)

```python
x = -3.2
y = 4.5

# Define variables r and θ (to define \theta and hit tab)
# Write your code here

```

```python
assert round(r - 5.5217750769114, 8) == 0
assert round(θ - 2.1889405527301, 8) == 0
```

```python
grader.check("q2")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: None)
```
>>> assert grade == grade_check
```
```
>>> assert grades == grades_check
```


Tests for q1 (points: None)
```
>>> assert even == [0, 2, 4, 6, 16]
```
```
>>> assert odd == [5, 7, 9, 11, 25, 39, 43]
```
```
>>> assert primes == [2, 5, 7, 11, 43]
```
```
>>> assert squares == [0, 9, 4, 25, 16]
```


Tests for q2 (points: None)
```
>>> assert round(r - 5.5217750769114, 8) == 0
>>> assert round(θ - 2.1889405527301, 8) == 0
```
