```python
# Initialize Otter
import otter
grader = otter.Notebook("00_variables.ipynb")
```

**Question 1:** Create new variables height and radius for a cylinder of height 0.75m and a
radius of 0.3 m. Write a Python code in Jupyter notebook to compute, store
and print the value of the following properties of a cylinder:<br>
(a) Surface area of the side<br>
(b) Surface area of each end<br>
(c) Total surface area (surface area of the sides + both ends)<br>
(d) Volume of the cylinder<br>
use `pi` value from the `math` library

```python
import ...
side_area = ...
end_area = ...
total_area = ...
volume = ...
print(side_area)
print(end_area)
print(total_area)
print(volume)
```

```python
grader.check("q1")
```

**Question 2:** A simply supported beam of length $2 m$ is supporting a uniformly distributed load $w$ of $500kN m^{-1}$. The Young's modulus of the beam $E$ is: $2\times10^{7} kPa$. The width of the beam $b$ is $0.3 m$ and the depth of the beam $d$ is $0.6 m$. Write a code creating appropriate variables to compute the deflection at any point $x$ on the beam. Calculate the deflection in $mm$ and the ratio of the current deflection to the maximum possible deflection at $x$ of 0.1, 0.25, 0.5, 1, and 1.7 m from the left support.

The deflection at any point $x$ on a simply supported beam with a UDL of $w$ is given as:
\begin{equation*}
    \delta = \frac{w x}{24 EI}(l^3 - 2lx^2 + x^3)
\end{equation*}

The maximum deflection of a simply supported beam with a UDL is:

\begin{equation*}
    \delta_{max} = \frac{5wl^4}{384EI}
\end{equation*}

The moment of inertia $I$ of a rectangular cross-section is given as:

\begin{equation}
    I = \frac{b d^3}{12}
\end{equation}

```python
#at x of 0.1m
deflection0 = ...
ratio0 = ...
#at x of 0.25m
deflection1 = ...
ratio1 = ...
#at x of 0.5m
deflection2 = ...
ratio2 = ...
#at x of 1m
deflection3 = ...
ratio3 = ...
#at x of 1.7m
deflection4 = ...
ratio4 = ...
print(f'max defelction: {max_deflection} mm\n')
print(f'deflection at x of 0.1m: {deflection0} mm')
print(f'ratio to the max deflection: {ratio0*100:.2f} %\n')
print(f'deflection at x of 0.25m: {deflection1} mm')
print(f'ratio to the max deflection: {ratio1*100:.2f} %\n')
print(f'deflection at x of 0.5m: {deflection2} mm')
print(f'ratio to the max deflection: {ratio2*100:.2f} %\n')
print(f'deflection at x of 1m: {deflection3} mm')
print(f'ratio to the max deflection: {ratio3*100:.2f} %\n')
print(f'deflection at x of 1.7m: {deflection4} mm')
print(f'ratio to the max deflection: {ratio4*100:.2f} %')
```

```python
grader.check("q2")
```

**Question 3:** Netflix employs 200 TB-sized (TerraBytes) HardDisks storage units at the local Internet Service Provider (ISP) to stream most frequently watched movies.
This reduces their bandwidth requirement and also makes watching the popular movies faster for their users. Consider the current storage capacity of
a Netflix storage device at your local ISP of 200TB, calculate how many HD
movies (approximately 1.5 GB) can be stored in that storage box? You have
been tasked by Netflix to procure new hard disks (storage) that can store as
many movies as the 200TB storage boxes could for HD movies, except Netflix
would like to store it in the new 4K format (approximately 6 GB) instead of
the HD format. What is the new storage size (GB) requirement?<br>
**Note : 1 TB is 1,000 GB**

```python
#number of HD movies stored in storage box
num_HD ...
#storage size (GB) for 4K movies
size_4k = ...
print(num_HD)
print(size_4k)
```

```python
grader.check("q3")
```

**Question 4:** The following Python code has bugs (errors in implementation) in the calculating of the following expression.

\begin{equation*}
(-4)^2 + \frac{208}{2 * 4}
\end{equation*}

Python expression: \-4**(2)+((208)/2*4)

(a) Fix the Python code so that it yields the correct answer of 42 for the above expression

(b) Remove any redundant (unnecessary) usage of parenthesis. 

(c) Enumerate the sequence of operations ($<object><operator><object>$) in which Python will execute the corrected code. (for e.g., Step 1: \-4**(2))

(d) What is the type (int/float/str) of the result from executing the above Python expression?


```python
correct_expression = ...
step1 = ...
step2 = ...
step3 = ...
step4 = ...
result_type = ...
print(correct_expression)
print(step1)
print(step2)
print(step3)
print(step4)
print(result_type)
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> round(side_area-1.413716694115407,5)==0
True
```
```
>>> round(end_area - 0.2827433388230814,5)==0
True
```
```
>>> round(total_area - 1.9792033717615696,5)==0
True
```
```
>>> round(volume - 0.21205750411731106,5)==0
True
```


Tests for q2 (points: None)
```
>>> round(deflection0-0.1535686728395062,5)==0
True
```
```
>>> round(ratio0-0.1592,2)==0
True
```
```
>>> round(deflection1-0.3744996624228396,5)==0
True
```
```
>>> round(ratio1-0.3883,2)==0
True
```
```
>>> round(deflection2-0.6872106481481483,5)==0
True
```
```
>>> round(ratio2-0.7125,2)==0
True
```
```
>>> round(deflection3-0.9645061728395063,5)==0
True
```
```
>>> round(ratio3-1.00,2)==0
True
```
```
>>> round(deflection4-0.44369212962963,5)==0
True
```
```
>>> round(ratio4-0.46,2)==0
True
```


Tests for q3 (points: None)
```
>>> num_HD==133333
True
```
```
>>> size_4k==799998
True
```


Tests for q4 (points: None)
```
>>> correct_expression == 42
True
```
```
>>> step1 == 16
True
```
```
>>> step2 == 8
True
```
```
>>> step3 == 26
True
```
```
>>> step4 == 42
True
```
```
>>> result_type is float
True
```
